---
title: chap '"<b>1234
description: chap '"<b>1234
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills:
  - 2
```

Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337

`@instructions`
Instruction "<b> 123 "<b> <img src=# onerror=alert(1); />133713371337

`@hint`
Hint '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337

`@pre_exercise_code`
```{python}
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337
```

`@sample_code`
```{python}
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337
```

`@solution`
```{python}
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337
```

`@sct`
```{python}

```
