---
title: chap '"<b>1234
description: chap '"<b>1234
---

## <script>alert(!)</script>aa

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills:
  - 2
```

Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337a

`@instructions`
Instruction "<b> 123 "<b> <img src=# onerror=alert(1); />133713371337

`@hint`
Hint '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337

`@pre_exercise_code`
```{python}
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337
```

`@sample_code`
```{python}
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337
```

`@solution`
```{python}
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337
```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: d9b3f46dde
xp: 100
```

<!-- Guidelines for contexts: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->![![![]()]()]()

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- Instruction 1
- Instruction 2

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
